<div class="alert alert-block alert-info">
<h3>Student Information</h3> Please provide information about yourself. We will NOT grade this submissing w/o all the information<br>
<b>Name</b>: Veton Abazovic<br>
<b>NetID</b>: va187<br>
<b>Recitation (01/02)</b>: 02<br>
<b>Notes to Grader</b> (optional):<br>
<br><br>
<b>IMPORTANT</b>
Your work will not be graded withour your initials below<br>
I certify that this lab represents my own work and I have read the RU academic intergrity policies at<br>
<a href="https://www.cs.rutgers.edu/academic-integrity/introduction">https://www.cs.rutgers.edu/academic-integrity/introduction </a><br>
<b>Initials</b>:      VA


<h3>Grader Notes</h3>
<b>Your Grade<b>:<br>
<b>Grader Initials</b>:<br>
<b>Grader Comments</b> (optional):<br>
</div>

## CS 439 - Introduction to Data Science
### Spring 2019

# Lab 3: Data Cleaning and Visualization

### Due Date: Sunday February 17, 2019 by 11:59 PM ###

### Instructions
This lab is presented as a notebook. Please execute the cells that are already completed and your task is to fill in the code
between ### BEGIN SOLUTION ### and ### END SOLUTION ###. 

#### Important: Please do not add any new cells or change the order of cells. If you have questions, please contact the courseS staff.

In this lab, you will be working with a dataset from NYPD containing data on calls to the New York Police Department. Information about the datasets can be found https://opendata.cityofnewyork.us/

## Setup

Note that after activating matplotlib to display figures inline via the IPython magic `%matplotlib inline`, we configure a custom default figure size. Virtually every default aspect of matplotlib [can be customized](https://matplotlib.org/users/customizing.html).

```python
import pandas as pd
import numpy as np
import zipfile
import matplotlib.pyplot as plt
import seaborn as sns

%matplotlib inline
plt.rcParams['figure.figsize'] = (12, 9)
```

# Part 1:Getting Data

We will work with the NYPD Historic complaint data set. Our first task is to estimate the size of this download by looking at the number of rows, columns and using an estimated size for a column (use a reasonable value). The site metadata is available from the page
https://data.cityofnewyork.us/Public-Safety/NYPD-Complaint-Data-Historic/qgea-i56i

```python
### BEGIN SOLUTION
estimated size of the download (based on metadata information from the site) =  2GB    
# please explain how you reached the answer
#The answer was given in class
### END SOLUTION
```

### 1.1 Download the data
This file is large (use the estimate you did above). If it takes too long to download, you may want to interrupt and download the file using a browser and URL https://data.cityofnewyork.us/Public-Safety/NYPD-Complaint-Data-Historic/qgea-i56i

```python
import utils
data_dir = 'data'
data_url = 'https://data.cityofnewyork.us/Public-Safety/NYPD-Complaint-Data-Historic/qgea-i56i'

file_name = 'NYPD_Complaint_Data_Historic.csv'

# To retrieve the dataset, we will use the `utils.fetch_and_cache` utility from utils library. 
dest_path = utils.fetch_and_cache(data_url=data_url, file=file_name,data_dir=data_dir)
print(f'Located at {dest_path}')
```

### 1.2 Inspect the size of the file
It is helpful to get an idea of the size of the file. This can be done using functions in the utils library.

```python
# look at the size of the file w/o opening it using OS (https://docs.python.org/3/library/os.html). You can perform
# variety of operating system related functions from this package.
### BEGIN SOLUTION
import os
statinfo = os.stat('data/NYPD_Complaint_Data_Historic.csv')
statinfo.st_size

#### END SOLUTION
```

### 1.3 Split the large file
This data file NYPD_Complaint_Data_Historic.csv is too big to load into a single DataFrame. Let us split the large file into smaller files.  Let us find out the number of lines in the NYPD_Complaint_Data_Historic.csv file using utils.

```python
# Using utils library, find the number of lines in the file

### BEGIN SOLUTION
num_of_line = utils.line_count('data/NYPD_Complaint_Data_Historic.csv')
num_of_line

### END SOLUTION
```

```python
# Split the file into 10 smaller files. Estimate the number of lines in each file using the count above
# files should be created in the data folder and named NYPD_Complaint_Data_Historic_1.csv, 
# NYPD_Complaint_Data_Historic_2.csv, ... NYPD_Complaint_Data_Historic_10.csv etc
# It is possible that few lines from the original file may not be saved due to rounding errors.

### BEGIN SOLUTION
def split(filehandler, delimiter=',', row_limit=603681,
          output_name_template='NYPD_Complaint_Data_Historic_%s.csv', output_path='./data', keep_headers=True):
    import csv
    reader = csv.reader(filehandler, delimiter=delimiter)
    current_piece = 1
    current_out_path = os.path.join(
        output_path,
        output_name_template % (current_piece)
    )
    current_out_writer = csv.writer(open(current_out_path, 'w', newline=''), delimiter=delimiter)
    current_limit = row_limit
    if keep_headers:
        headers = next(reader)
        current_out_writer.writerow(headers)
    for i, row in enumerate(reader):
        if i + 1 > current_limit:
            current_piece += 1
            current_limit = row_limit * current_piece
            current_out_path = os.path.join(
                output_path,
                output_name_template % (current_piece)
            )
            current_out_writer = csv.writer(open(current_out_path, 'w', newline=''), delimiter=delimiter)
            if keep_headers:
                current_out_writer.writerow(headers)
        current_out_writer.writerow(row)

split(open('data/NYPD_Complaint_Data_Historic.csv', 'r'));
### END SOLUTION
```

### 1.4 Check the files in the data folder
Now, we'll use a method of the `Pathlib.Path` class called `glob` to list all files in the `data` directory. You will find useful information in pathlib [docs](https://docs.python.org/3/library/pathlib.html).

Below, we use pathlib's `glob` method to store the list of all files' names from the `data_dir` directory in the variable `file_names`. These names should be strings that contain only the file name (e.g. `dummy.txt` not `data/dummy.txt`). The asterisk (*) character is used with the `glob` method to match any string.

```python
from pathlib import Path
data_dir_path = Path('data') # creates a Path object that points to the data directory
file_names = [x.name for x in data_dir_path.glob('*') if x.is_file()]
file_names
```

### 1.5 Pre-processing of data
It is good to pre-process the data to see if the file can be opened in a Jupyter notebook. We need to avoid large files that can crash notebooks. Typically, files of size around 200 MB is ok to open into a DataFrame. In the following activities we will inspect the file w/o opening it as a DataFrame. Using utils.head 

```python
# Use the util.head() to read up to 5 lines from the original file (w/o opening it)
from utils import head
head('data/NYPD_Complaint_Data_Historic.csv')
```

### 1.6 Inspecting and describing data columns
There should be 35 columns in each record. Using header information and data types, describe the type of data in each
column. If you are unable to determine, just state so.
##### BEGIN SOLUTION
* CMPLNT_NUM : Nominal
* CMPLNT_FR_DT: Ordinal
* CMPLNT_FR_TM:Ordinal
* CMPLNT_TO_DT:Ordinal
* CMPLNT_TO:ordinal
* ADDR_PCT_CD: ordinal
* RPT_DT: ordinal
* 
* ....
* VIC_RACE:Nominal
* VIC_SEX:Nominal
##### END SOLUTION

# Part 2 - Exploratory Data Analysis
Exploratory data analysis (EDA) is the process of examining a subest of a large data set to see what we can know about the data. First we will explore one file NYPD_Complaint_Data_Historic_1.csv to see what we can find out.

### 2.1 Loading Data into a DataFrame
Load the first CSV file, NYPD_Complaint_Data_Historic_1.csv into a `pandas.DataFrame` object. Also do a time analysis to see how long it took to load the data into a DataFrame. Time should be printed in seconds. The time libraries https://docs.python.org/3/library/time.html can help.

```python
### BEGIN SOLUTION
import pandas as pd
import time
start = time.clock()
calls = pd.read_csv("data/NYPD_Complaint_Data_Historic_1.csv")
elapsed = (time.clock() - start)
elapsed


### END SOLUTION
```

### 2.2 Description of Fields
Let's also check some basic information about these files using the `DataFrame.describe` and `DataFrame.info` methods. Describe columns that can be removed based on the information.

```python
### BEGIN SOLUTION
calls.info()
calls.describe()

# What columns can be removed from the DataFrame? A reasonable rule of thumb is that if a column is missing more than 
# 50% of the data then it should be removed
#PARKS_NM
#STATION_NAME
#HADEVELOPT         
#HOUSING_PSA
### END SOLUTION
```

### 2.2 Finding Uniques

Notice that the functions above reveal type information for the columns, as well as some basic statistics about the numerical columns found in the DataFrame. However, we still need more information about what each column represents. Let's explore the data further.

find the number of unique values in each DataFrame column and answer the questions below.

```python
### BEGIN SOLUTION
calls.nunique()

# Questions
# 1. How many distinct locations where the complaints have come from?77081
# 2. How many age groups are represented in the data set? 16
# 3. How many boroughs are included in the data set? 5
# 4. How many offense types are listed in this data set? 67


### END SOLUTION
```

### 2.3 Offense by Boro
Using GroupBy operation, create a DataFrame that groups offenses by Boro. call the DataFrame calls_by_Boro_and_offense

<!-- grade_id: q1a-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
calls_by_Boro_and_offense = calls.groupby(['BORO_NM', 'OFNS_DESC'])['OFNS_DESC'].agg('count')
calls_by_Boro_and_offense
### END SOLUTION
```

<!-- grade_id: q1a-tests, points: 1, grade: true, solution: false -->
```python
## Assert that in Bronx there were 6022 Grand Larceny incidents
### BEGIN SOLUTION
assert calls_by_Boro_and_offense["BRONX", "GRAND LARCENY"] == 7815
### END SOLUTION
```

### 2.4 Offenses in Bronx

In the cell below, find a list of strings corresponding to the possible values for `OFNS_DESC` when `BORO` is "BRONX". Create an expression that automatically extracts the names of the offenses.

<!-- grade_id: q1b-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df = calls_by_Boro_and_offense['BRONX'].reset_index(name = 'OFNS')
answer24 = df['OFNS_DESC'].tolist()
print(answer24)
### END SOLUTION

# How many offenses were committed in Bronx during the analysis period?
offense_made = calls_by_Boro_and_offense['BRONX'].size
offense_made
```

### 2.5 Most Common Crimes in NYC

What are the five crime types of OFNS_DESC that have the most crime events in Bronx? You may need to use `value_counts` to find the answer. Save your results as a list of strings.

**Hint:** *The `keys` method of the Series class might be useful.*

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
temp = df.sort_values(ascending=False, by='OFNS').head()
answer25 = temp['OFNS_DESC'].tolist()
answer25
### END SOLUTION
```

### 2.6 Probability of a Crime in Bronx
What is the probability that a the crime "Arson" can happen in Bronx?

<!-- grade_id: q2-tests, points: 1, grade: true, solution: false -->
```python
### BEGIN SOLUTION
df1 = df[df['OFNS_DESC'] == 'ARSON']
answer26 = (df1['OFNS'].sum()/df['OFNS'].sum())
answer26
### END SOLUTION
```

---
# Part 3: Visualizing the Data

## Pandas vs. Seaborn Plotting

Pandas offers basic functionality for plotting. For example, the `DataFrame` and `Series` classes both have a `plot` method. However, the basic plots generated by pandas are not particularly pretty. While it's possible to manually use matplotlib commands to make pandas plots look better, we'll instead use a high level plotting library called Seaborn that will take care of most of this for us.

As you learn to do data visualization, you may find the [pandas documentation](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.plot.html) and [Seaborn documentation](https://seaborn.pydata.org/api.html) helpful!

### 3.1 Plotting a Series
Using the built-in plotting functionality of pandas, such as `plot` method of the `Series` class to generate a `barh` plot type,  display the value counts for `OFNS_DESC` visually as a barh chart.

```python
### BEGIN SOLUTION
ofns=calls["OFNS_DESC"].value_counts()
ofns.plot('barh')
plt.xlabel("Number of Calls")
plt.ylabel("Crime Category")
plt.title("Number of Calls by Crime Type")
### END SOLUTION
```

### 3.2 Getting a Better Plot
The plot above can be messy as it plots all offenses. Plot only the offenses that has more than 10000 calls

```python
### BEGIN SOLUTION
ofns = calls['OFNS_DESC'].value_counts()
s = pd.Series(ofns).where(lambda x : x>10000).dropna()
s.plot('barh')
plt.xlabel('Number of Calls')
plt.ylabel('Crime Category')
plt.title('Number of Calls by Crime Type')
### END SOLUTION
```

By contrast, the Seaborn library provides a specific function `countplot` built for plotting counts. It operates directly on the DataFrame itself i.e. there's no need to call `value_counts()` at all. This higher level approach makes it easier to work with. Use the y-label ("Crime Category"), x-label("Number of Calls") and title_of_plot("Number of Calls By Crime Type"

```python
### BEGIN SOLUTION
ax = sns.countplot(y = 'OFNS_DESC', data = calls, orient = "h")
ax.set_xlabel('Number of Calls')
ax.set_ylabel('Crime Category')
ax.set_title("Number of Calls By Crime Type")
### END SOLUTION
```

You may note that the ordering might be different for the seaborn plot (as compared to pandas plot). If we want the same ordering that we had in the pandas plot, we can use the order parameter of the `countplot` method. It takes a list of strings corresponding to the axis to be ordered. By passing the index of the `value_counts`, you can get the order you want.

```python
### BEGIN SOLUTION
ax = sns.countplot(y = 'OFNS_DESC', data = calls, order = calls['OFNS_DESC'].value_counts(ascending = True).index, orient = "h")
ax.set_xlabel('Number of Calls')
ax.set_ylabel('Crime Category')
ax.set_title("Number of Calls By Crime Type")

### END SOLUTION
```

Now we have a pretty bar plot with the bars ordered by size. Though `seaborn` appears to provide a superior plot from a aesthetic point of view, the `pandas` plotting library is also good to understand. You'll get practice using both libraries in the following questions.

## An Additional Note on Plotting in Jupyter Notebooks

You may have noticed that many of our code cells involving plotting end with a semicolon (;). This prevents any extra output from the last line of the cell that we may not want to see. Try adding this to your own code in the following questions!

### 3.3 making more plots

Now it is your turn to make some more plots using `pandas` and `seaborn`. Let's start by looking at the distribution of calls over days of the week.

The CMPLNT_FR_DT field contains the date of the event. We would like to add a new column to the DataFrame that includes Day of the week (DAY_OF_WEEK) that indicates the day of the week. This can help us analyze the crimes on a specific day of the week. For example, we can answer questions such as "what day of the week that a LARSON is likely to happen in NYC?"


Add a new column `DAY_OF_WEEK` into the `calls` dataframe that has the day string (eg. 'Sunday') for the corresponding value in CMPLNT_FR_DT. For example, if the first 3 values of `CMPLNT_FR_DT` are `['01/27/2006, '01/28/2006, '01/29/2006]`, then the first 3 values of the `DAY_OF_WEEK` column should be `["Friday", "Saturday", "Sunday"]`.

**Hint:** *Try using the [Series.map](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.map.html) function on `calls["OFNS_DESC"]`.  Can you assign this to the new column `calls["DAY_OF_WEEK"]`?*

<!-- grade_id: q3a-answer, points: 0, grade: false, solution: true -->
```python
days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"]
### BEGIN SOLUTION
temps1 = calls
temps1['CMPLNT_FR_DT'] = pd.to_datetime(temps1['CMPLNT_FR_DT'])
temps1['DAY_OF_WEEK'] = temps1['CMPLNT_FR_DT'].dt.weekday_name
temps1
### END SOLUTION
```

### 3.4 Seaborn plots

Create a `seaborn` plot that shows the number of calls for each day of the week. You may want to use of the `rotation` argument in `ax.set_xticklabels`, which rotates the labels by 90 degrees.

```python
### BEGIN SOLUTION
ax = sns.countplot(x='DAY_OF_WEEK', data = temps1, order= days)
plt.xlabel('DAY_OF_WEEK')
plt.ylabel('Number of Calls')
plt.title('Number of Calls For Each Day of the Week')
ax.set_xticklabels(days, rotation=90)


### END SOLUTION
```

Now, let's make the same plot using `pandas`. Construct a vertical bar plot with the count of the number of calls (entries in the table) for each day of the week **ordered by the day of the week** (eg. `Sunday`, `Monday`, ...). Do not use `sns` for this plot. Be sure that your axes are labeled and that your plot is titled.

**Hint:** *Given a series `s`, and an array `coolIndex` that has the same entries as in `s.index`, `s[coolIndex]` will return a copy of the series in the same order as `coolIndex`.*

<!-- grade_id: q3b-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
weekdf= temps1.groupby('DAY_OF_WEEK')['DAY_OF_WEEK'].agg('count').reset_index(name = 'WEEKDAY').reindex([3,1,5,6,4,0,2])
ax = weekdf.set_index('DAY_OF_WEEK').plot(kind = 'bar', title = 'Number of Calls For Each Day of the Week', legend = False)
#plt.xticks(range(len(days)), days)
ax.set_xlabel('DAY_OF_WEEK')
ax.set_ylabel('Number of Calls')
### END SOLUTION
```

## 3.5 What Day of the Week is more calls?

Is it true that weekdays generally have slightly more calls than Saturday or Sunday? What can you say about the difference?

##### BEGIN SOLUTION
The weekdays usually have more calls than Saturday or Sunday as observed from the graphs above. Fridays tend to have the most calls then Wednesday has the second most while Saturday comes in Third. Sunday is almost tied with Monday. There is a slight differences.

##### END SOLUTION

We can break down into some particular types of events to see their distribution. For example, let's make a bar plot for the OFNS_DESC "HARRASSMENT 2". Which day is the peak for "HARRASSMENT 2"?

This time, use `seaborn` to create a vertical bar plot of the number of total noise violations reported on each day of the week, again ordered by the days of the week starting with Sunday. Do not use `pandas` to plot.

**Hint:** *If you're stuck, use the code for the seaborn plot in above question as a starting point.*

<!-- grade_id: q4a-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
harrass = calls[calls['OFNS_DESC'] == 'HARRASSMENT 2']
ax = sns.countplot(x='DAY_OF_WEEK', data = harrass, order = days)
plt.xlabel('DAY_OF_WEEK')
plt.ylabel('Number of Calls')
plt.title('Number of HARRASSMENT 2 Calls For Each Day of the Week')
ax.set_xticklabels(days, rotation=90)
### END SOLUTION
```

### 3.6 

Do you see anything interesting about the distribution of HARRASSMENT 2 calls over a week? Type a short answer below.
##### BEGIN SOLUTION
The interesting thing about the distribution of Harrassment 2 calls over a week is that there is more calls during the weekday than weekend.
##### END SOLUTION

### 3.7 More Plots

Let's look at a similar distribution but for a crime we have much more calls data about. In the cell below, create the same plot as you did in previous questions, but now looking at instances of the OFNS_DESC "BURGLARY" (instead of "HARRASSMENT 2"). Use either `pandas` or `seaborn` plotting as you desire.

<!-- grade_id: q5-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
burg = calls[calls['OFNS_DESC'] == 'BURGLARY']
ax = sns.countplot(x='DAY_OF_WEEK', data = burg, order = days)
plt.xlabel('DAY_OF_WEEK')
plt.ylabel('Number of Calls')
plt.title('Number of BURGLARY Calls For Each Day of the Week')
ax.set_xticklabels(days, rotation=90)

### END SOLUTION


```

### 3.8 time of events

Now let's look at the CMPLNT_TO_TM column which indicates the time for events. Since it contains hour and minute information, let's extract the hour info and create a new column named `Hour` in the `calls` dataframe. You should save the hour as an `int`. Then plot the frequency of each hour in the table (i.e., `value_counts()`) sorted by the hour of the day (i.e., `sort_index()`).

You will want to look into how to use:

* [Series.str.slice](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.str.slice.html#pandas.Series.str.slice) to select the substring.
* [Series.astype](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.astype.html) to change the type.

**Hint:** *The `str` helper member of a series can be used to grab substrings.  For example, `calls["CMPLNT_TO_TM"].str.slice(3,5)` returns the minute of each hour of the `CMPLNT_TO_TM`.*

<!-- grade_id: q6a-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
timedf = calls
timedf['HOUR'] = ((timedf['CMPLNT_TO_TM'].str.slice(0,2)).str.replace(':',''))
#timedf['HOUR'].str.replace(':','').fillna(0)
pd.to_numeric(timedf['HOUR'])
ax = sns.countplot(y = 'HOUR', data = timedf, order = timedf['HOUR'].value_counts(ascending=True).index)

### END SOLUTION
```

Create a pandas bar plot showing the number of BURGLARY crimes committed at each hour of the day. Use the labels
* ax.set_xlabel("Hour of the Day")
* ax.set_ylabel("Number of Calls")
* ax.set_title("Number of Calls Reporting Fraud For Each Day of the Week");

```python
### BEGIN SOLUTION
plot = timedf['HOUR'].value_counts()
t = timedf['HOUR'].value_counts().sort_index(axis=0).plot(kind="bar")
t.set_ylabel('Hour of the Day')
t.set_xlabel('Number of Calls')
plt.title('Number of Calls Reporting Fraud for Each Day of Week')

### END SOLUTION
```

### 3.9 More plots

In the cell below, create a seaborn plot of the same data. Again, make sure you provide axes labels and a title for your plot.

<!-- grade_id: q6b-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ordr = [0,1,2,3,4,5,6,7,8,9,10,11,12,13,14,15,16,17,18,19,20,21,22,23]
ax = sns.countplot(x='HOUR', data=timedf)
### END SOLUTION
```

### 3.8 Spike in burglary?

According to your plots, is there a spike in calls reporting BURGLARY at any particular time? If so, Do you trust that this spike is legitimate, or could there be an issue with our data? Explain your reasoning in 1-2 sentences below.

#### BEGIN SOLUTION
I feel that the data is incorrect as there should not be a spike in the morning while people are getting up and its usually busy in the morning.

#### END SOLUTION

In the cell below, we generate a boxplot which examines the hour of day of each crime broken down by the `OFNS_DESC` value.  To construct this plot we used the [DataFrame.boxplot](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.boxplot.html) documentation.

```python
timedf.boxplot(column='HOUR', by='OFNS_DESC', rot=90);
```

While the pandas boxplot is informative, we can use seaborn to create a more visually-appealing plot. Using seaborn, regenerate a better box plot. See either the textbook (https://www.textbook.ds100.org/ch/06/viz_quantitative.html) or the [seaborn boxplot documentation](https://seaborn.pydata.org/generated/seaborn.boxplot.html).

Looking at your plot, which crime type appears to have the largest interquartile range? Put your results into `answer` as a string.

<!-- grade_id: q7-answer, points: 0, grade: false, solution: true -->
```python
# Todo: Make a boxplot with seaborn
### BEGIN SOLUTION
ax = sns.boxplot(x='OFNS_DESC', y='HOUR', data = timedf)
ax.set_xticklabels(ax.get_xticklabels(),rotation=90)
ax.set_title("Distributions of Calls Over Hours in a Day, Grouped by Crime");
answer = "Your Answer here"
### END SOLUTION
```

### 3.9 - Visualization of crimes on a Map of NYC
finally we attempt to visualize the crimes committed in NYC on a Map. First we need to installing some mapping software. run the cell below to install folium package for mapping software.

```python
!pip install --upgrade folium
```

```python
### Plot the maps for BURGLARY in NYC
### if it takes too much time or map does not show up, try plotting a subset of the data set of for a specific crime

### BEGIN SOLUTION


### END SOLUTION
```

## Congratulations !!!

<div class="alert alert-block alert-info">
<h2>Submission Instructions</h2> 
<b> File Name:</b> Please name the file as your_section_your_netID_Lab3.jpynb<br>
<b> Submit To: </b> Canvas &rarr; Assignments &rarr; Lab3 <br>
<b>Warning:</b> Failure to follow directions may result in loss of points.<br>
</div>

Credits: Josh Hug, and Berkeley Data Science Group for their contributions to the original version.